<img style="float: left;" src="earth-lab-logo-rgb.png" width="150" height="150" />

# Earth Analytics Education - Bootcamp Course Fall 2020

## Important  - Assignment Guidelines

1. Before you turn in your assignment, make sure to run the entire notebook with a fresh kernel. To do this first, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart & Run All)
2. In the cells below you will replace the `raise NotImplementedError()` code with your code that addresses the activity challenge. If you don't replace that code, your notebook will not run properly.

```
# YOUR CODE HERE
raise NotImplementedError()
```

3. Any open ended questions will have a "YOUR ANSWER HERE" within a markdown cell. Replace that text with your answer also formatted using Markdown.
4. **IMPORTANT: DO NOT RENAME THIS NOTEBOOK!** If the file name changes, the autograder will not grade your assignment properly.
5. **Do not rename the notebook file.** If you do, the autograder will not recognize your submisson.
6. When you plot, please comment out `plt.show()` as the code below will effectively run `plt.show()` for you and also will grab your plot for autograding. DO NOT DELETE any code that says `DO NOT REMOVE LINE BELOW`. That code is for autograding!!

```
### DO NOT REMOVE LINE BELOW ###
student_plot1_ax = nb.convert_axes(plt)
```



## Follow to PEP 8 Syntax Guidelines

* Run the `autopep8` tool on all cells prior to submitting (HINT: hit shift + the tool to run it on all cells at once!
* Use clear and expressive names for variables. 
* Organize your code to support readability.
* Check for code line length
* Use comments and white space sparingly where it is needed


### Add Your Name Below 
**Your Name:**

<img style="float: left;" src="colored-bar.png"/>

---

# Week 8 Homework - Loops & Conditionals

To complete assignment 10, review Chaps 17-18 in Section 7 of the <a href="https://www.earthdatascience.org/courses/intro-to-earth-data-science/" target="_blank">Intro to Earth Data Science online textbook</a> online textbook, which conditional statements and loops in **Python**.   

Read the instructions for each question carefully to successfully complete the required tasks.


## Assignment Data

For this assignment, you will write **Python** code to download and work with data on fire occurrence in California between 1992 and 2015 provided by <a href="https://www.fs.usda.gov/rds/archive/Product/RDS-2013-0009.4/" target="_blank">the United States Forest Service</a>.

```python
# DO NOT MODIFY THIS CELL
# Core imports needed for grading
import matplotcheck.notebook as nb
from matplotcheck.base import PlotTester
import matplotcheck.autograde as ag
import matplotlib
```

## Import Python Packages

In the cell below, add code **after the line for `Your Code Here`**, replacing `raise NotImplementedError()` with your code to import the packages/modules needed to:
* create plots 
* set your working directory
* get lists of files in a directory
* download data using earthpy functions
* work with pandas DataFrames

Be sure to list the package imports following the appropriate PEP 8 order and spacing. 


<!-- grade_id: student-imports-answer, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Set Working Directory

In the cell below complete the following task:

* **Use a conditional statement** to:
    * Set the working directory to **your `earth-analytics/data` directory in your home directory** if the path exists.
    * Make the directory using `os.makedirs()` and then print a helpful message that the path does not exist but it's being created. 

Reduce repetition in your code with reusable variables. 

Be sure to use the appropriate functions that will allow your code to run successfully on any operating system.

<!-- grade_id: student-set-wd-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Download Data Using EarthPy

In the cell below, add code **after the line for `Your Code Here`**, replacing `raise NotImplementedError()` with your code, to complete the following task:

* **Use EarthPy to download this url: https://ndownloader.figshare.com/files/25033508**: 

In this download, you'll see three directories with files containing data by year:

* ca-fires-yearly/monthly-fire-count
    * The dataset contains the total number of fires (greater than 100 acres) that occurred in each month and year in California between 1992 and 2015. The data are organized with a file for each year from 1992 to 2015.
* ca-fires-yearly/monthly-mean-size
    * The dataset contains the mean fire size (acres) of all fires greater than 100 acres for each month and year in California between 1992 and 2015. The data are organized with a file for each year from 1992 to 2015
* ca-fires-yearly/1992-2015-gt-100-acres
    * The dataset contains the cause of all fires greater than 100 acres in California between 1992 and 2015. The data are organized with a file for each year from 1992 to 2015

<!-- grade_id: student-download-data-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: student-download-tests, points: 5, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
# Tests that the working directory is set to earth-analytics/data

path = os.path.normpath(os.getcwd())
student_wd_parts = path.split(os.sep)

wd_points = 0

if student_wd_parts[-2:] == ['earth-analytics', 'data']:
    print("\u2705 Great - it looks like your working directory is set correctly to .../earth-analytics/data")
    wd_points += 5
else:
    print("\u274C Oops, the autograder will not run unless your working directory is set to earth-analytics/data")

print("\n \u27A1 You received {} out of 5 points for setting your working directory.".format(
    wd_points))
wd_points
```

<img style="float: left;" src="colored-bar.png"/>

## Challenge 1: Import Monthly Fire Count Data

In the cell below complete the following task:

* Create a for loop that prints the file name  for each `.csv` file in the `earthpy-downloads\ca-fires-yearly\monthly-fire-count` directory
* In that same loop, extract the date for each year in each `.csv` file name and add it to a new list sorted in order from earliest to most recent  year using `list.append()`.

Your final list should look something like this (with more years however): 

`["1992", "1993", "1994"]`

At the end of the cell, call the list object of years so that it prints below 
the cell as standard output (**do not use the print statement**).

<!-- grade_id: student-import-count-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: student-import-count-tests-1, points: 0, grade: true, solution: false -->
```python
# Test that the list was created and returned correctly

student_year_list = _

if isinstance(student_year_list, list):
    print("\u2705 The object you created is a list, good job!")
else:
    print("\u274C The object you created is not a list, please make sure to create a list.")
```

<!-- grade_id: student-import-count-tests-2, points: 5, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
```

## Challenge 2: Create A Monthly Fire Count DataFrame With a Year Column

In the cell below use the same for loop structure that you used above to create a dataframe containing 
data for all of the years. The dataframe should have month names as a header and should contain monthly fire counts for each year. 

1. loop through each `.csv` file sorted by name
2. open that file using pandas
3. add a new column to each data.frame that contains the year associated with the data. You will need to extract the year from the file name similar to what you did above. 
3. add the DataFrame object to a list object using `list-name.append()`
4. When you loop is complete:
   * combine all of the dataframes in the list into a new `DataFrame` using `pd.concat(list-name, ax=0)` to combine all of the dataframes in your list into a final dataframe that contains all of the data.
   * Set the "year"  column as an index

Call the final `DataFrame` object at the end of the cell. 

<!-- grade_id: student-count-year-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: student-count-year-test-1, points: 0, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
student_dataframe_monthly_count = _

if isinstance(student_dataframe_monthly_count, pd.DataFrame):
    print("\u2705 Object created is a dataframe, good job!")
else:
    print("\u274C Object created is not a dataframe.")

if student_dataframe_monthly_count.index.name == 'year':
    print("\u2705 Successfully read in the year column as the index!")
else:
    print("\u274C The index is not set to the year column.")
```

<!-- grade_id: cell-student-count-year-test-2, points: 5, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
```

## Challenge 3: Calculate Monthly Mean Fire Count

In the cell below use Pandas to calculate the **mean monthly fire count**.
You should have one value for each month as an output. 

Set the display options to:

`pd.options.display.float_format = '{:.1f}'.format`

To ensure that the values returned are rounded.
Be sure to call the object at the end of the cell so that it renders as cell output.
**DO NOT USE THE PRINT STATEMENT!**

<!-- grade_id: student-monthly-mean-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: student-monthly-mean-test, points: 5, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
```

<img style="float: left;" src="colored-bar.png"/>

## Challenge 4: Calculate Yearly Mean Fire Count

In the cell below use pandas to calculate the mean yearly fire count.
HINT: similar to `numpy.mean()`, using `axis=` when you call `.mean()` will help you 
calculate a mean per row (by year) rather than by column (by month). 

For your answer below, use `pd.options.display.float_format = '{:.1f}'.format` to 
ensure that your data values are rounded to one decimal point. 

IMPORTANT: Call the object that contains mean values at the end of the cell so the 
output renders as output below.

<!-- grade_id: student_yearly_mean_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: student_yearly_mean_test, points: 5, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
```

<img style="float: left;" src="colored-bar.png"/>

## Challenge 5: Open and Create a Fire Size DataFrame With A Year Column

Use a `for` loop to open and add the data in the directory:

`earthpy-downloads/ca-fire-data/yearly_data/mean_fire_size`

to a new DataFrame following the same steps that you completed above.

1. loop through each `.csv` file sorted by name
2. open that file using pandas
3. add a new column to each `DataFrame` that contains the year associated with the data. You will need to extract the year from the file name similar to what you did above. 
4. Outside of the for loop, combine all of the files into a single DataFrame and set the year column to be an index (just like you did for the fire count data above). 

Call the new DataFrame object at the end of the cell.

<!-- grade_id: create-fire-size-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: create-fire-size-test-1, points: 0, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
# Visible test that there is a DataFrame with a year index

student_size_df = _

if isinstance(student_size_df, pd.DataFrame):
    print("\u2705 Object created is a dataframe, good job!")
else:
    print("\u274C Object created is not a dataframe.")


if student_size_df.index.name == 'year':
    print("\u2705 Successfully read in the year column as the index!")
else:
    print("\u274C The index is not set to the year column.")
```

<!-- grade_id: create-fire-size-test-2, points: 5, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
```

## Challenge 6: Calculate the Monthly  Max Values for Fire Size 

In the cell below calculate the max fire size value **for each month** (summarized across all years).

You should have one value per month as the output. 
Be sure to call the object in the cell so it renders below (DO NOT USE PRINT).

<!-- grade_id: fire-size-monthly-max-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: fire-size-monthly-max-test-1, points: 0, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
student_monthly_max_size = _

if len(student_monthly_max_size) == 12:
    print("\u2705 Great work - you have 12 objects in your {} object".format(type(student_size_df.max())))
else:
    print("\u274C Oops! You don't have 12 objects in your output.")
```

<!-- grade_id: fire-size-monthly-max-test-2, points: 5, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
```

<img style="float: left;" src="colored-bar.png"/>

## Challenge 7: Figure of Fire Season in California

You have been asked to create a figure that describes the number of fires (fire count)
and the associated size of fires in California. The goal of this figure it so 
look at what time of the year we are seeing the largest and most fires. 

In the cell below complete the following task:

* Create one multi-plot figure that contains two subplots that are **vertically spaced - one on top 
of the other**:
    * **top plot**: Plot the mean **fire count** for each month
    * **bottom plot** : Plot the mean of **fire size** for each month
* **Use a different color for each plot but you can use the same style if you like.** 
    * For each plot, be sure to include appropriate titles and axes labels including units of measurement where appropriate. 
* Add an overall title for the entire figure. 

For your title and labels, be sure to think about the following pieces of information that could help someone easily interpret the plot:
* geographic coverage or extent of data.
* duration or temporal extent of the data.
* what was actually measured and/or represented by the data.
* units of measurement.

**REMINDER:** comment out the code line `plt.show()` in your plot code like this: `# plt.show()`.

<!-- grade_id: student-create-plot-1-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

### DO NOT REMOVE LINE BELOW ###
hist_plot_1 = nb.convert_axes(plt, which_axes="all")
```

<!-- grade_id: plot-1-tests, points: 20, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
count_hist = PlotTester(hist_plot_1[0])
size_hist = PlotTester(hist_plot_1[1])
bar_plot_1, bar_plot_2 = False, False

try:
    count_hist.assert_plot_type('bar')
    print("\u2705 First plot is a bar plot!")
    bar_plot_1 = True
except AssertionError:
    print("\u274C The first plot is not a bar plot, make sure to make it to a bar plot. Tests will not run successfully if this does not pass.")

try:
    size_hist.assert_plot_type('bar')
    print("\u2705 Second plot is a bar plot!")
    bar_plot_2 = True
except AssertionError:
    print("\u274C The second plot is not a bar plot, make sure to make it to a bar plot. Tests will not run successfully if this does not pass.")

```

## Discuss Your Plot

In the Markdown cell below, answer the following questions using a Markdown list.

1. In what month are fires the largest during the time period represented in the two datasets?
2. In what month are there the most fires during the time period represented in this data?
3. If the above answers are different months, can you think of a reason why this trend might occur?


<!-- grade_id: cell-e5b21fde84e3e84e, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

<img style="float: left;" src="colored-bar.png"/>

## Challenge 8: Fires by Cause Data

In the cell below do the following: 

1. Loop through, open and combine all of the files in the 
`earthpy-downloads/ca-fires-yearly/1992-2015-gt-100-acres` directory
2. Add a year column to the data
3. Calculate the total count of fires by "cause" using pandas `df.size()`.

<!-- grade_id: open-cause-data-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: open-cause-data-test-1, points: 0, grade: true, solution: false -->
```python
student_fires_by_cause = _

if isinstance(student_fires_by_cause, pd.Series):
    print("\u2705 Your output is a Pandas Series, good job!")
else:
    print("\u274C Your output is not a Pandas Series, make sure you are running groupby at the end of the cell.")
```

<!-- grade_id: open-cause-data-test-2, points: 8, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
```

In the cell below subset your fires DataFrame so it only contains fires
that are of fire_class_size "G".

<!-- grade_id: filter-g-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: filter-g-test, points: 4, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
# Test class g filter

student_class_g_fires = _

class_g_points = 0

if len(student_class_g_fires == 394):
    print("\u2705 Correct number of rows in your class G dataframe!")
    class_g_points += 2
else:
    print("\u274C There are not the correct amount of rows in your class G dataframe.")

if all(x == 'G' for x in student_class_g_fires.fire_size_class.to_list()):
    print("\u2705 Successfully filtered dataframe to only contain class G fires!")
    class_g_points += 2
else:
    print("\u274C Dataframe does not contain only class G fires.")

print("\n \u27A1 You received {} out of 4 points for filtering the class G fires correctly.".format(
    class_g_points))
class_g_points
```

## Challenge 9: Subset Fire Cause Data

In the cell below, create two dataframe objects:

1. Create one that  contains data for the years 1995 to 2000 (this DF should contain values for 1995 through 2000).
2. Create a dataframe object that contains data for the years1995 to 2000 (this DF should contain values for 2010 through 2015).

At the end of the cell, call both objects on the same line using the syntax:

`df-one, df-two`

The first and last 5 rows for both objects should then print below the cell.

<!-- grade_id: student-class-g-time-filter-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: class-g-filter-tests-2, points: 0, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
# Tests to make sure your dataframes were correctly made

student_g_fires_by_year = _

if isinstance(student_g_fires_by_year[0], pd.core.frame.DataFrame):
    print("\u2705 Great, your first object is a Pandas DataFrame!")
else:
    print("\u274C Didn't call both dataframes at the end of the cell.")

if isinstance(student_g_fires_by_year[1], pd.core.frame.DataFrame):
    print("\u2705 Great, your second object is a Pandas DataFrame!")
else:
    print("\u274C Didn't call both dataframes at the end of the cell.")

student_df_lengths = len(student_g_fires_by_year[0]), len(
    student_g_fires_by_year[1])

if (84, 88) == student_df_lengths or (88, 84) == student_df_lengths:
    print("\u2705 DataFrames have the correct number of rows in them!")
else:
    print("\u274C One or both of the dataframes doesn't have the correct number or rows in the results.")
```

<img style="float: left;" src="colored-bar.png"/>

## Challenge 10: Figure - Fires by Cause

Create a figure with two stacked subplots that rendered as **scatterplots** (points) using `ax.plot()` to plot the `1992-2015-gt-100-acres` data. 

For each subplot: 

* Only include points that represent fires of size_class `G` (this is data that you subsetted above).
* Color the points in your plot using the `cause` column in your data. There are a few ways to do this, but for this exercise, color the points by `human-caused` vs `Non-human caused` using a **conditional statement**
within the for loop that you write to create your plot. 

The following causes should be labeled as `human-caused`

`["Arson", "Smoking", "Equipment Use", "Campfire", "Powerline", "Railroad"]`

(HINT: To make the plot look like a scatter plot you can adjust the line and marker options on the plot)

You will use the data / `DataFrame` that you worked with above in the  `1992-2015-gt-100-acres` directory to produce this plot. 

1. The top subplot will contain fires colored by cause for the years: 2010-2015
2. The bottom subplot will contain fires colored by cause for the years 1995-2000


To create these plots, you will need to create a for loop that plots groups pandas dataframes. 

*********
HINT:

* A conditional statement can be used to test whether a string is in an existing list of strings:
```
list_of_things = ["item1", "item2"]

if acategory in list_of_things:
    # Set color
else:
    # Set color  
```

For your plot, be sure to do the following:

1. Use a shared x and y axis by adding `sharex` and `sharey` to your `plt.subplots()` call
example: `plt.subplots(sharex=True)`
2. Use `.scatter()` to create your plot.
3. Add titles to each subplot
4. Color each category by human vs other cause. You should have a total of two colors  on your plot.
5. Add an overall figure title
6. Add a legend for each plot. The legend should have multiple "causes" (e.g. arson, lightning, etc) but only two different point colors (one for human caused and one for non human caused) 


<!-- grade_id: fires-by-cause-figure-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()


### DO NOT REMOVE LINE BELOW ###
fire_causes_plot = nb.convert_axes(plt, which_axes="all")
```

<!-- grade_id: fire-cause-by-figure-test-1, points: 0, grade: true, solution: false -->
```python
# NOTE: 2010-2015 Data NEEDS TO BE THE TOP PLOT FOR THE AUTOGRADER TO WORK!

if len(fire_causes_plot) == 2:
    print("\u2705 You successfully added two plots to your figure, good job!")
else:
    print("\u274C There are not two plots in your figure, make sure you plot the 2000 and 2015 data.")
```

<!-- grade_id: hidden-tests-answer, points: 0, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
```

<!-- grade_id: fire-cause-by-figure-test-2, points: 32, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
```

## Challenge 11: Discuss Your Plot

Look at the plot above. Do you notice a difference in fire size between 1995-2000 and 2010-2015?
Add your answer to the markdown cell below.

<!-- grade_id: cell-4c52b3438af0547d, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Pep 8, Spelling and Does the Notebook Run?
In this cell, we will give you points for the following

1. PEP 8 is followed throughout the notebook (4 points)
2. Spelling and grammar are considered in your written responses above (4 points)
3. The notebook runs from top to bottom without any editing (it is reproducible) - 4 points